---
title: 'Environment  pour la réalisation de TP noté'
description: 'Cet environnement contient déjà tout les packages nécessaire pour tester les commandes qui vous seront demandées dans le cadre du TP noté.'
---

## Environment pour tester votre code

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

A l'aide du package ... charger les données

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
