---
title: 'Environment  pour la réalisation de TP noté'
description: 'Cet environnement contient déjà tout les packages nécessaire pour tester les commandes qui vous seront demandées dans le cadre du TP noté.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

A l'aide du package ... charger les données

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}
# Charger la librarie ggplot
library(readr) 
deces_mensuels = "https://assets.datacamp.com/production/repositories/4170/datasets/28a0c9bc5b08dd9f3db3ee93e4c5e6be4ea8609e/monthly_deaths.csv"
par_clinique = "https://assets.datacamp.com/production/repositories/4170/datasets/d14eae8c81d12f2ea8d4a1433942b2b449d90fc6/yearly_deaths_by_clinic.csv"
read_csv(par_clinique)
```

`@sample_code`
```{r}
# Charger la librarie ggplot
monthly_deaths.csv
```

`@solution`
```{r}

```

`@sct`
```{r}

```
